---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Introduction to Convolutional Neural Networks

```yaml
type: VideoExercise
key: 55a7c3bb45
xp: 50
```

`@projector_key`
5b71d68181d2c17f7b150f2e56187580
